---
title: "An instructor's introduction to learnr"
author: "Monica Thieu, PhD (almost)"
institute: "Dept of Psychology, Columbia University"
date: "updated: `r Sys.Date()`"
output:
  xaringan::moon_reader:
    css: xaringan-themer.css
    lib_dir: libs
    nature:
      highlightStyle: github
      highlightLines: true
      countIncrementalSlides: false
---

```{r setup, include=FALSE}
options(htmltools.dir.version = FALSE)
# Assume most chunks are only to be shown as examples and not actually run
knitr::opts_chunk$set(eval = FALSE)
require(tidyverse)
```

```{r xaringan-themer, eval=TRUE, include=FALSE, warning=FALSE}
library(xaringanthemer)
style_mono_light(
  base_color = "#0046a6",
  header_font_google = google_font("Josefin Sans"),
  text_font_google   = google_font("Montserrat", "300", "300i"),
  code_font_google   = google_font("Fira Mono")
)
```

## Learning goals

By the end of this presentation, you will be able to:

--

- Describe the pros and cons of learnr for instruction

--

- Evaluate the fit of your R lesson to a learnr tutorial

--

- Break down your R lesson into conceptual learnr-compatible chunks

--

This is not a tutorial on learnr! The [learnr package docs](https://rstudio.github.io/learnr/articles/learnr.html) and [RStudio forums](https://community.rstudio.com/tag/learnr) are excellent resources for actually setting up a tutorial.

--

However, I will be highlighting specific learnr features or syntax when they are relevant to conceptual assignment design.

---

class: inverse, middle

# What is learnr?

---

## What is learnr?

.pull-left[

learnr is a package that allows you to publish documents with **interactive, runnable R code chunks.**

You can use it for many applications, including **interactive textbook-style self-teaching documents** and **problem sets.**

]

.pull-right[

![learnr logo](https://rstudio.github.io/learnr/logo.png)

]

---

layout: true

## How does learnr work?

---

Very high-level:

--

`learnr` uses the Shiny prerendered R Markdown document engine to create interactive documents.

--

After downloading the package, configuring it as a tutorial only requires setting the R Markdown YAML header:

```yaml
title: "My learnr tutorial"
output: learnr::tutorial
runtime: shiny_prerendered
```

---

Chunks are turned into interactive exercises using the `exercise = TRUE` chunk option.

> Change the code below to calculate 1 + 2 instead.

```{r}
# {r question-1, exercise = TRUE}
1 + 1
```

--

`learnr` documents differ from regular Rmds in one key way:

For the most part, each exercise chunk must be able to run independently. **Exercise chunks cannot inherit the output of prior exercise chunks.**

???

https://stackoverflow.com/questions/70024831/chunks-in-text-rmarkdown-xaringan

---

However, you can pre-run setup code for a code chunk.

--

For example, you can pre-`library()` a package:

```{r}
# {r setup, include = FALSE}
library(tidyverse)
```

--

Or you can preload a dataframe, so that it is accessible to students at the beginning of the exercise:

```{r}
# {r preload-data}
load("data/example_data.rda")
```

```{r}
# {r question-2, exercise = TRUE, exercise.setup = "preload-data"}
example_data
```

---

You can pre-fill exercises, so students only need to add or modify code.

> Fill in the correct argument/value to calculate the mean of the `price` column in `diamonds`, excluding any NA values.

```{r}
# {r question-4, exercise = TRUE}
mean(diamonds$price, ___ = ___)
```

---

.pull-left[

You can include non-exercise code chunks and hide their source with `echo = FALSE`, so that students can see the solution output and compare their output to it.

> Add the correct geom to produce a scatter plot of diamond price vs. diamond weight in carats.

]

.pull-right[

```{r, eval = TRUE, dev = "svglite", fig.asp = 1/2}
# {r question-5-solution}
ggplot(diamonds,
       aes(x = carat,
           y = price)) +
  geom_point()
```

```{r}
# {r question-5, exercise = TRUE}
ggplot(diamonds,
       aes(x = carat,
           y = price))
```

]

---

.pull-left[

You can auto-grade students' responses via `learnr`'s integration with the `gradethis` package.

It isn't all-powerful, but it can save you time and help students self-teach while doing assignments!

]

.pull-right[

![gradethis package logo](https://pkgs.rstudio.com/gradethis/logo.png)

]

---

layout: false

## When to just have students fill in an R Markdown?

`learnr` is not always the best use of your efforts, though, depending on the nature of your assignment.

```{r, eval = TRUE, echo = FALSE}
tribble(~Criterion, ~learnr, ~Rmd,
                "Assess whether students can write:", "code snippets", "a full analysis",
                "Students need to produce:", "code & multiple choice answers", "code & short answer text",
                "You want to grade:", "programmatically", "visually",
                "Do students need to be able to knit?", "no", "yes") %>% 
  knitr::kable()
```


---
class: inverse, middle

# Designing a learnr assignment

---

## How to design a learnr-compatible problem set

Approximately following the principles of [backward design](https://cft.vanderbilt.edu/guides-sub-pages/understanding-by-design/)...
--

- Define the right answer

--

- Modularize exercises

--

- Pre-fill exercise chunks with intention

---

layout: true

### Define the right answer

---

--

**Operationalize what counts as the correct answer.**

--

Decide what the output _object_ needs to look like. A single value, a vector, a dataframe, a model object--it can be any R object!

--

Ideally, you will be able to test that the student's output is equivalent to a solution object.

--

You can do this several ways, but the most straightforward way is to write solution code, and `learnr` will automatically compare the result of the student's code with the result of the solution chunk.

```{r}
# {r question, exercise = TRUE}
diamonds
```

```{r}
# {r question-solution}
diamonds %>% 
  group_by(color) %>% 
  summarize(mean_carat = mean(carat))
```

---

Decide what _functions_ need to be used to produce that output. 

--

You can also grade an exercise based on the functions used.

--

For example, including `gradethis::grade_code()` in another chunk along with a solution chunk will check the written code is equivalent, not just the result.

```{r}
# {r question-code-check}
gradethis::grade_code()
```

---

Identify common wrong answers, so that you can give students more informative hints if they make a common mistake.

--

You can customize some grading functions to give specific messages when the result matches some particular type of wrong answer.

```{r}
# {r question-check}
gradethis::grade_result(
  fail_if(~identical(.result,
                     diamonds %>% summarize(mean_carat = mean(carat))),
          "You forgot to group the data before summarizing."),
  pass()
)
```

---

layout: true

### Modularize exercises

---

--

Once you have defined the right answer (and some predictable wrong ones) for a particular exercise...

--

**Break down your assignment into manageable, grade-able code snippets.**

--

This makes it easier for you to scaffold students' code. If you're teaching in the tidyverse, this logic plays very well with the sequential nature of the pipe (`%>%` or `|>`).

--

It also makes it easier to grade the result of smaller pieces of code (fewer commands means fewer possible wrong answers).

--

_However, remember that you cannot make an exercise chunk dependent on students' responses in prior chunks!_

--

You can deal with this in one of two ways.

---

**Option 1:** You can "spoil" the right answer to a previous exercise in the current exercise.

--

> Use `group_by()` to group the `diamonds` dataframe by diamond color category.

```{r}
# {r question, exercise = TRUE}
diamonds %>% 
  group_by(color)
```

--

> Next, use `summarize()` on the grouped data to calculate the mean weight in carats of each diamond color category, in a new column called `mean_carat`.

```{r}
# {r question, exercise = TRUE}
diamonds %>% 
  group_by(color) %>% 
  summarize(mean_carat = mean(carat))
```

---

You can require the previous exercise to be answered correctly in order to reveal the current exercise, using progressive reveal:

--

```yaml
title: "My learnr tutorial"
output:
  learnr::tutorial:
**    progressive: true
runtime: shiny_prerendered
```

--

However, this requires that students be able to see whether their exercise code is "correct" upon running the chunk.

--

By default, when code solutions/grading are provided, `learnr` allows students to see whether their code is correct or incorrect when they submit, but you can turn this off.

--

You may have reasons for not wanting students to be able to immediately see whether they've gotten the right answer, in which case this option will not work well for you.

---

**Option 2:** You can also just ask students to fill in a multi-line answer in a single exercise.

--

> Use `group_by()` and `summarize()` to calculate the mean weight in carats of each diamond color category, in a new column called `mean_carat`.

```{r}
# {r question, exercise = TRUE}
diamonds
```

--

There more ways for students to get the wrong answer with the increased flexibility, so you may get more panicky/confused emails about these questions. Still, this may be the right amount of scaffolding (not that much) depending on the topic and the students.

---

layout: true

### Pre-fill exercise chunks with intention

---

--

Once you have your solutions defined and chunks modularized, you can decide what default code shows up for students when they start each exercise.

--

You can leave exercise chunks totally empty, or you can fill them in with some amount of helper code.

`learnr` allows you to put blanks in place of function & variable names, using three underscores `___`. The exercise will prompt students to replace the blanks with real code.

---

> Use `group_by()` and `summarize()` to calculate the mean weight in carats of each diamond color category, in a new column called `mean_carat`.

--

Least amount of helper code:

```{r}
# {r question, exercise = TRUE}
diamonds
```

--

Slightly more helper code:

```{r}
# {r question, exercise = TRUE}
diamonds %>% 
  ___() %>% 
  ___()
```

--

Even more helper code:

```{r}
# {r question, exercise = TRUE}
diamonds %>% 
  ___() %>% 
  ___(mean_carat = ___)
```

---

You want to balance the following goals when deciding how much helper code to write:

--

- Giving away the right amount of answer

--

- Giving enough hints/structure so students know where to start

--

- Helping with things they don't need to have learned yet

--

- Pre-filling things like column names that you need spelled correctly for programmatic grading
